Mono-methylated histones control PARP-1 in chromatin and transcription

Author response:

The following is the authors’ response to the previous reviews

We extend our sincere gratitude for the invaluable comments provided by the reviewers and yourself, along with the constructive suggestions to enhance the quality of our manuscript. In response to this invaluable feedback, we have diligently revised and resubmitted our paper as an article, introducing five primary figures, seven supplementary figures, and two supplementary data files. Importantly, this work represents a significant contribution to the field, presenting novel findings for the first time without any prior publication.

Within the enclosed document, we have provided a comprehensive response to the editor and reviewer comments, addressing each point meticulously and specifically. We extend our heartfelt thanks to the reviewers and yourself for your diligent examination of our manuscript and for offering insightful recommendations.

In our latest revision, we have taken great care to address every comment, ensuring that we clarify the manuscript and provide robust evidence where required. We have meticulously highlighted the modifications within the manuscript in yellow for your convenience, while also including the modifications made in response to each specific comment. The primary focus of these revisions was to provide additional context regarding the relationship between PARP-1 and mono-methylated histones. Substantial modifications were made to our discussion section to address this point.

Another concern raised was regarding the discrepancy in the relationship of PR-SET7 and PARP-1 between our study and the recent study by Estève et al. (PMID: 36434141). We have revised the results and discussion sections to discuss this concern.

Addressing Reviewer 2’s concern about the potential indirect role of PARP1 in the regulation of some metabolic genes despite its direct binding to loci coding for metabolic genes we revised the discussion section to highlight this possibility.

Enclosed, you will find a detailed, point-by-point response to each of the editor’s and reviewers' comments, showcasing our commitment to addressing their concerns with precision.

We firmly believe that our revisions successfully resolve all the concerns raised by the editor and the reviewers, and we are confident that this improved version of our manuscript contributes significantly to the scientific discourse. Once again, we thank you for considering our work, and please feel free to contact me if you require any additional information.

In the revised manuscript, most of the concerns raised by the reviewers have been addressed satisfactorily. However, as suggested by reviewer#2, it would have been more significant, if the PARP1-mediated reading of global mono-methylation of histone could be addressed. At least the mechanisms of selectivity of PARP1 need further convincing discussion.

We thank the editor for their valuable comments. We have extended our discussion section to discuss in more detail the relationship between PARP1 and mono-methylated histones. In our refined Discussion section, we have endeavored to articulate more clearly how PARP-1 may be selectively recruited to active chromatin domains through its interaction with mono-methylated histone marks. We propose a model wherein PARP-1 actively participates in the turnover process, contributing to the maintenance of an active chromatin environment. This mechanism entails PARP-1 selectively binding to mono-methylated active histone marks associated with highly transcribed genes. Upon activation, PARP-1 undergoes automodification, leading to its release from chromatin and facilitating the reassembly of nucleosomes carrying the mono-methylated marks. Subsequently, the enzymatic action of Poly(ADP)-ribose glycohydrolase (PARG) cleaves pADPr, enabling the restoration of PARP-1's binding affinity to mono-methylated active histone marks. This proposed hypothesis is consistent with existing research across various model organisms and aligns with the known association of PARP-1 with highly expressed genes, as well as its role in mediating nucleosome dynamics and assembly.

Our modified Discussion section unfolds as follows:

"Finally, highly transcribed genes have been reported to present a high turnover of mono-methylated modifications, maintaining a state of low methylation (50). Moreover, our previous study revealed that PARP1 preferentially binds to highly active genes (34). Consequently, our findings suggest an active involvement of PARP-1 in the turnover process to maintain an active chromatin environment. This proposed mechanism unfolds in the following steps: 1) PARP-1 selectively binds to mono-methylated active histone marks associated with highly transcribed genes. 2) Upon activation, PARP-1 undergoes automodification and subsequently disengages from chromatin, facilitating the reassembly of nucleosomes carrying the mono-methylated marks. 3) The enzymatic action of Poly(ADP)-ribose glycohydrolase (PARG) cleaves pADPr, restoring PARP-1's binding affinity to mono-methylated active histone marks. This proposed hypothesis is consistent with existing research conducted across various model organisms, including mice, Drosophila, and Humans (7, 24, 30, 51-53). Notably, previous studies have consistently demonstrated that PARP-1 predominantly associates with highly expressed genes and plays a crucial role in mediating nucleosome dynamics and assembly. Thus, our proposed model provides a molecular framework that may contribute to understanding the relationship between PARP-1 and the epigenetic regulation of gene expression."

We trust that these revisions effectively address the editor’s comment and enhance the overall strength and clarity of our manuscript.

Furthermore, recent developments in the area are omitted, as an important publication hasn't been discussed anywhere in the work (PMID: 36434141).

We appreciate the editor's thorough review of our revised manuscript and the responses to the previous reviewer's comments. To address this important concern, we have carefully investigated the levels of PR-SET7 in parp1 hypomorphic conditions.

Supplemental Fig. S4 and S5 demonstrate that in the absence of Parp1, there were no significant changes observed in PR-SET7 RNA or protein levels, respectively. This finding supports the conclusion that Parp1 is not directly involved in the regulation of PR-SET7 in Drosophila contrasting with the findings of Estève et al.'s study (PMID: 36434141). This discrepancy may arise from differing relationships between PARP-1 and PR-SET7, which could cooperate in the context of Drosophila development while playing antagonistic roles in specific cell lines or under particular conditions.

We have updated the Results section to explicitly mention this observation:

"Interestingly, in the absence of PARP-1, neither PR-SET7 RNA nor protein levels were affected (Supplemental Fig.S4-5), indicating that PARP-1 is not directly implicated in the regulation of pr-set7. This finding contrasts with recent evidence demonstrating PARP1-induced degradation of PR-SET7/SET8 in human cells (16)."

Furthermore, we have modified the discussion section to address this discrepancy:

"A recent study demonstrated that in human cells overexpressing PARP-1, PR-SET7/SET8 is degraded, whereas depletion of PARP-1 leads to an increase in PR-SET7/SET8 levels (16). However, in our study involving parp-1 mutant in Drosophila third-instar larvae revealed a nuanced scenario: we detected a minor but not significant reduction in both PR-SET7 RNA and protein levels (Supplemental Fig.S4 and S5). This outcome stands in stark contrast to the previous study's findings. The discrepancy could be due to the distinct experimental approaches used: the previous research focused on mammalian cells and in vitro experiments, whereas our study examined the functions of PARP-1 in whole Drosophila third-instar larvae during development. Consequently, while PARP-1 may cooperate with PR-SET7 in the context of Drosophila development, it could exhibit antagonistic roles against PR-SET7 in specific cell lines and under certain biological or developmental conditions."

We believe that these modifications effectively address the raised concern and provide a more comprehensive understanding of the relationship between PARP1 and PR-SET7 in our study. We hope these clarifications enhance the overall robustness and clarity of our findings.

Reviewer #2 (Public Review):

Summary:

This study from Bamgbose et al. identifies a new and important interaction between H4K20me and Parp1 that regulates inducible genes during development and heat stress. The authors present convincing experiments that form a mostly complete manuscript that significantly contributes to our understanding of how Parp1 associates with target genes to regulate their expression.

Strengths:

The authors present 3 compelling experiments to support the interaction between Parp1 and H4K20me, including:

(1) PR-Set7 mutants remove all K4K20me and phenocopy Parp mutant developmental arrest and defective heat shock protein induction.

(2) PR-Set7 mutants have dramatically reduced Parp1 association with chromatin and reduced poly-ADP ribosylation.

(3) Parp1 directly binds H4K20me in vitro.

Weaknesses:

(1) The RNAseq analysis of Parp1/PR-Set7 mutants is reasonable, but there is a caveat to the author's conclusion (Line 251): "our results indicate H4K20me1 may be required for PARP-1 binding to preferentially repress metabolic genes and activate genes involved in neuron development at co-enriched genes." An alternative possibility is that many of the gene expression changes are indirect consequences of altered development induced by Parp1 or PR-Set7 mutants. For example, Parp1 could activate a transcription factor that represses metabolic genes. The authors counter this model by stating that Parp1 directly binds to "repressed" metabolic genes. While this argument supports their model, it does not rule out the competing indirect transcription factor model. Therefore, they should still mention the competing model as a possibility.

We appreciate Reviewer 2's insightful comments during both rounds of revision, which have significantly enriched the quality of our manuscript. The binding of PARP1 to loci encoding metabolic genes indeed suggests a direct role of PARP1 in their regulation. However, we acknowledge Reviewer 2's point that some of these targets might be regulated indirectly, with PARP1 potentially modulating the expression of intermediary transcription factors.

To address this possibility, we have revised the discussion section of our manuscript accordingly:

"Remarkably, our observations indicate a notable affinity of PARP-1 for binding to the gene bodies of these metabolic genes (34), suggesting a direct involvement of PARP1 in their regulation. Nonetheless, it remains plausible that certain genes may be indirectly regulated by PARP1 through intermediary transcription factors."

We trust that this modification adequately addresses Reviewer 2's concern.

(2) The section on inducibility of heat shock genes is interesting but missing an important control that might significantly alter the author's conclusions. Hsp23 and Hsp83 (group B genes) are transcribed without heat shock, which likely explains why they have H4K20me without heat shock. The authors made the reasonable hypothesis that this H4K20me would recruit Parp-1 upon heat shock (line 270). However, they observed a decrease of H4K20me upon heat shock, which led them to conclude that "H4K20me may not be necessary for Parp1 binding/activation" (line 275). However, their RNA expression data (Fig4A) argues that both Parp1 and H40K20me are important for activation. An alternative possibility is that group B genes indeed recruit Parp1 (through H4K20me) upon heat shock, but then Parp1 promotes H3/H4 dissociation from group B genes. If Parp1 depletes H4, it will also deplete H4K20me1. To address this possibility, the authors should also do a ChIP for total H4 and plot both the raw signal of H4K20me1 and total H4 as well as the ratio of these signals. The authors could also note that Group A genes may similarly recruit Parp1 and deplete H3/H4 but with different kinetics than Group B genes because their basal state lacks H4K20me/Parp1. To test this possibility, the authors could measure Parp association, H4K20methylation, and H4 depletion at more time points after heat shock at both classes of genes.

We sincerely appreciate Reviewer 2 for their insightful comment on our manuscript. Your hypothesis regarding the potential induction of H3/H4 dissociation from group B genes by PARP-1, leading to a reduction in H4K20me1, offers a thought-provoking perspective. However, our findings suggest an alternative interpretation.

Our data indicate that while H4K20me1 is indeed present under normal conditions at group B genes, its reduction following heat shock does not seem to impede PARP-1's role in transcriptional activation (Fig. 4A, C, and E). Instead, we propose that this decrease in H4K20me1 might signify a regulatory shift in chromatin structure, facilitating transcriptional activation during heat shock, with PARP-1 playing an independent facilitating role. Moreover, existing studies have highlighted the dual role of H4K20me1, acting as a promoter of transcription elongation in certain contexts and as a repressor in others.

The elevated enrichment of H4K20me1 in group B genes under normal conditions may indeed indicate a repressive state that requires alleviation for transcriptional activation. Additionally, we cannot discount the possibility of unique regulatory functions associated with PR-SET7, extending beyond its recognized role as a histone methylase. Non-catalytic activities and potential interactions with non-histone substrates might contribute to the nuanced control exerted by PR-SET7 on group B genes during heat stress.

Furthermore, our exploration of pr-set720 and ParpC03256 mutants reveals distinct roles for PARP-1 and H4K20me1 in modulating gene expression (Fig 3E). This reinforces the notion that the interplay between PR-SET7 and PARP-1 involves a multifaceted regulatory mechanism.

To address these points, we have revised the discussion section of our manuscript accordingly:

"Another plausible explanation could be that the recruitment of PARP-1 to group B genes loci promotes H4 dissociation and then leads to a reduction of H4K20me1. However, our findings suggest an alternative interpretation: the decrease in H4K20me1 at group B genes during heat shock does not seem to impede PARP-1's role in transcriptional activation, (Fig.4A, C and E). Rather than disrupting PARP-1 function, we propose that this reduction in H4K20me1 may signify a regulatory shift in chromatin structure, priming these genes for transcriptional activation during heat shock, with PARP-1 playing an independent facilitating role. Moreover, existing studies have highlighted the dual role of H4K20me1, acting as a promoter of transcription elongation in certain contexts and as a repressor in others (13, 26, 39, 40, 42-46). The elevated enrichment of H4K20me1 in group B genes under normal conditions may indicate a repressive state that requires alleviation for transcriptional activation. Additionally, we cannot discount the possibility of unique regulatory functions associated with PR-SET7, extending beyond its recognized role as a histone methylase. Non-catalytic activities and potential interactions with non-histone substrates might contribute to the nuanced control exerted by PR-SET7 on group B genes during heat stress (47, 48). Furthermore, our exploration of pr-set720 and parp-1C03256 mutants reveals distinct roles for PARP-1 and H4K20me1 in modulating gene expression (Fig 3E). This reinforces the notion that the interplay between PR-SET7 and PARP-1 involves a multifaceted regulatory mechanism. Understanding the intricate relationship between these molecular players is crucial for elucidating the complexities of gene expression modulation under heat stress conditions."

We believe that this modification enhances the clarity of our conclusions and adequately addresses Reviewer 2's concerns regarding the intricate relationship between PARP-1, H4K20me1, and PR-SET7 in transcriptional regulation under heat stress conditions.

Changes from the earlier version to the revised version

Abstract

Abstract

PARP-1 is central to transcriptional regulation under both normal and stress conditions, with the governing mechanisms yet to be fully understood. Our biochemical and ChIP-seq-based analyses showed that PARP-1 binds specifically to active histone marks, particularly H4K20me1. We found that H4K20me1 plays a critical role in facilitating PARP-1 binding and the regulation of PARP-1-depenednt loci during both development and heat shock stress. Here we report that the sole H4K20 mono-methylase, pr-set7, and parp-1 Drosophila mutants undergo developmental arrest. RNA-seq analysis showed an absolute correlation between PR-SET7- and PARP-1-dependent loci expression, confirming co-regulation during developmental phases. PARP-1 and PR-SET7 are both essential for activating hsp70 and other heat shock genes during heat stress, with a notable increase of H4K20me1 at their gene body. Mutating pr-set7 disrupts monometylation of H4K20 along heat shock loci and abolish PARP-1 binding there. These data strongly suggest that H4 monometylation is a key triggering point in PARP-1 dependent processes in chromatin.

Introduction

The chromosomes of eukaryotes are made up of chromatin, which is a complex of DNA tightly wound around nucleosomes (1). Nucleosomes are the basic unit of chromatin, and they are made up of about 147bp of DNA wrapped around an octamer of four core histones (H2A, H2B, H3 and H4) (1). Eukaryotic DNA is packed into nucleosomes to maintain higher order chromatin structure. As a consequence, DNA sequences are inaccessible to a plethora of protein effectors for biological processes such as DNA damage repair, replication and transcription (2). To overcome the nucleosome barrier, eukaryotes have developed several mechanisms that ensure the successful removal of nucleosomes. One of the key effectors involved in chromatin remodeling is Poly(ADP)-ribose Polymerase 1 (PARP-1), a chromatin-associated nuclear enzyme implicated in a host of biological processes including, transcriptional regulation and DNA repair (3–5). PARP-1 catalyzes the addition of poly(ADP)-ribose (pADPr) moieties onto acceptor proteins by utilizing donor NAD+ as substrate. Upon PARP-1 activation due to developmental and environmental cues or stress, PARP-1 is automodified, and it attaches negatively charged poly(ADP)-ribose polymers onto histones and other chromatin-associated proteins, thereby decondensing DNA via electro-repulsion and facilitating PARP-1 mediated biological processes (6, 7). Besides PARP-1, histone modifications are essential for remodeling chromatin structure. Histones are subject to post-translational modifications, including phosphorylation, acetylation, sumoylation, ubiquitylation and ADP-ribosylation on their amino-terminal tails protruding from the nucleosome core (2, 8, 9). Modifications of these histone tails can alter chromatin structure by modifying inter-nucleosome interactions or by serving as docking sites for effectors to facilitate diverse biological outcomes depending on the histone(s) so modified (2, 10).

The binding of PARP-1 to post-translationally modified histones may be necessary for PARP-1-mediated biological processes and PARP-1 targeting to chromatin. In our previous study, we showed that the phosphorylation of the histone variant H2Av directly impacts the localization, activation, and enzymatic activity of PARP-1 (11). Also, we previously showed that the N-terminal tails of histones H3 and H4 are needed for the enzymatic activation of PARP-1 (12). Taken together, there is strong evidence that histone modifications can regulate PARP-1 activity in vivo.

In this study, we demonstrate that PARP-1’s transcriptional regulatory role under regular and heat stress conditions may be intricately tied to its interaction with specific mono-methylated active histone marks, notably H4K20me1, H3K4me1, H3K36me1, H3K27me1 and H3K9me1 that it is inhibited by repressive H3K9me2/3 marks. Through transcriptomic analysis of pr-set7 and parp-1 mutant Drosophila third-instar larvae, we uncover a high correlation in differentially expressed genes and co-enrichment of PARP-1 and H4K20me1 at a subset of these genes, underscoring the role of PR-SET7/H4K20me1 in facilitating PARP-1-mediated transcriptional regulation. Importantly, we show that PARP-1 and PR-SET7 are essential for the activation of specific heat shock genes including hsp70, and the dynamic regulation of H4K20me1 during this process, further validates our hypothesis of H4K20me1 being crucial for PARP-1 binding and gene regulation under developmental and heat stress conditions.

Results

PARP-1 binds to H4K20me1, H3K4me1, H3K36me1, H3K9me1 and H3K27me1 in vitro and in vivo

In this study, we investigated the interaction between full-length recombinant PARP-1 and various histone modifications using a histone peptide array containing modified histone peptides individually and in combination with other peptides (Fig.1A). Our previous work established that PARP-1 binds to core histones (12). Here, we aimed to determine the specific histone marks that modulate PARP-1’s affinity for chromatin. We found that PARP-1 binds specifically to spots containing H4K20me1, H3K4me1, H3K36me1, H3K9me1 or H3K27me1 peptides (Fig.1B, C, D and data file S1). Specificity analysis showed that PARP-1 binding was highly specific for H3K4me1, H3K36me1, H3K9me1 and H3K27me1 containing peptides with the highest specificity observed for H4K20me1 (Fig.1E and data file S1). Additionally, we explored the inhibitory effects of different histone modifications on PARP-1 binding. Our findings showed that H3K9me2/3 hindered PARP-1 binding to spots containing H3K4me1 (Fig.1F, Supplemental Fig.S1 and S2). Conversely, H3K4me1 enhanced PARP-1 binding to spots with H3K9me2/3 (Fig.1F and Supplemental Fig.S1). In parallel, our data shows that the phosphorylation of H3S10 or H3T11, two residues nearby H3K9 diminish PARP-1 binding to H3K9me1 (Supplemental Fig.S1 and S2). Similarly, the phosphorylation of H3K28 affects PARP-1 binding to H3K27me1 (Supplemental Fig.S1 and S2).

PARP-1 binds H4K20me1, H3K4me1, H3K36me1, H3K9me1, and H3K27me1 in vitro and in vivo.

(A) Illustration of the MODified histone peptide assay (Active Motif) used to determine PARP-1 binding to histone modifications. The histone peptide array (Top, left), comprising 19mer peptides with single or up to four concurrent histone modifications, was employed to investigate PARP-1’s binding affinity for histone modifications and to assess the impact of adjacent modified peptides on PARP-1 binding. This array was first blocked, then incubated with PARP-1 protein (Top, right). Subsequently, it was stained with a PARP-1 antibody and a horseradish peroxidase (HRP)-conjugated secondary antibody (Bottom, right). Visualization of PARP-1 binding was done through enhanced chemiluminescence detection and captured on X-ray film (Bottom, left). See Methods for a full description. (B) Signal intensity on modified histone peptide array based on incubation with PARP-1 protein. (C) Average intensities of PARP-1 binding to single histone peptides. (D-E) Reproducibility and specificity of spot intensities from modified histone peptide array duplicates. (D) Scatter plot showing the correlation of the average intensities of duplicate arrays. Intensities of PARP-1 binding to all peptides and spots containing single H4K20me1, H3K4me1, H3K36me1, H3K9me1 and H3K27me1 (key) are shown. Pearson’s correlation coefficient (r) is 0.92. (E) Bar chart showing top 8 histone modifications with the highest specificity for PARP-1 binding. The specificity factor was calculated by dividing the average intensity of spots that contain the modified histone peptide by the average intensity of spots that do not contain the peptide. (F) Spearman correlation of PARP-1, H4K20me1, H3K4me1, H3K36me1, H3K9me1, H3K9me2/3 and H3K27me1 peaks in Drosophila third-instar larvae based on fraction of overlap. (G) Heatmaps showing k-means clustering-generated occupancy of PARP-1, H4K20me1, H3K4me1, H3K36me1, H3K9me1, H3K9me2/3 and H3K27me1 normalized ChIP-seq signals in third-instar larvae at Drosophila genes. ChIP-seq signals are sorted in descending order. The upper plots show the summary of the signals.

Next, we examined the association of PARP-1 with H4K20me1, H3K4me1, H3K36me1, H3K9me1, H3K9me2/3 and H3K27me1 in Drosophila third-instar larvae using ChIP-seq data. PARP-1 peaks were highly correlated with H4K20me1, H3K4me1, H3K36me1, H3K9me1 and H3K27me1 but significantly less correlated with H3K9me2/3 peaks (Fig.1G). Consistently, at gene clusters, H4K20me1 and H3K4me1 were highly enriched at gene bodies in cluster I which also had the highest PARP-1 enrichment at promoters. H3K36me1, H3K9me1 and H3K27me1 were also enriched in cluster I albeit to a lesser degree than H4K20me1 and H3K4me1 (Fig.1H). In contrast, H3K9me2/3 were more enriched at clusters II and III but depleted in cluster I (Fig.1H). Additionally, PARP-1 associates with H4K20me1 and H3K4me1 on Drosophila polytene chromosomes (Fig.2). Notably, H4K20me1, H3K4me1, H3K36me1, H3K9me1 and H3K27me1 are associated with active genes while H3K9me2/3 is associated with low-expression or silent genes (13). Thus, our data shows that PARP-1 predominantly binds to active mono-methylated histone marks, specifically H4K20me1, H3K4me1, H3K36me1, H3K9me1 and H3K27me1 while also suggesting that the binding of PARP-1 may be hindered by repressive histone marks such as H3K9me2/3.

PARP-1 colocalizes with H3K4me1 and H4K20me1 histone marks in polytene chromosomes.

(A) Immunofluorescent staining of Drosophila salivary gland chromosomes showing colocalization of PARP-1 with H3K4me1 and H4K20me1. White lines indicate areas of the colocalization-quantification shown in (B). (B-C) Quantification of fluorescence intensity of PARP-1, H3K4me1 and H4K20me1 at Drosophila polytene chromosomes in panel A. (B) Images show distribution of PARP-1 fluorescence intensity with H3K4me1 (top) and H4K20me1 (bottom) fluorescence intensity. (C) Images represent a scatterplot showing PARP-1 fluorescence intensity with H3K4me1 (top) and H4K20me1 (bottom) fluorescence intensity. Pearson correlation coefficients (r) are respectively 0.58 (top) and 0.71 (bottom).

PARP-1 and PR-SET7 co-regulate developmental gene expression programs

Earlier version: Next, we examined the relationship between H4K20me1 and PARP-1 during transcriptional regulation. Given that flies that lack PR-SET7, the sole methylase of H4K20, and PARP-1 undergo developmental arrest during the larval-to-pupal transition (14, 15), we hypothesized that PR-SET7/H4K20me1 and PARP-1 may regulate similar gene expression programs during development. To validate our hypothesis, we initially confirmed that the pr-set720 mutant not only eliminated PR-SET7 RNA and protein but also abrogated H4K20me1 modification (Supplemental Fig.S3). Interestingly, in the absence of PARP-1, neither PR-SET7 RNA nor protein levels were affected (Supplemental Fig.S4-5), indicating that PARP-1 is not directly implicated in the regulation of pr-set7. Subsequently, we integrated PARP-1 and H4K20me1 ChIP-seq data in third-instar larvae and RNA-seq of pr-set720 and parp-1C03256mutant third-instar larvae. Our data showed that the differentially expressed genes (DEGs) in pr-set720 and parp-1C03256 transcriptome were highly correlated (Pearson, r = 0.79) (Fig.3A and B). However, despite the global reduction of H4K20me1, and PARP-1 RNA in pr-set720 and parp-1C03256 mutants, respectively (14, 15), there was a limited alteration in gene expression at genes where PARP-1 was co-localized with H4K20me1 (Fig.3C and D). Differentially expressed genes in both pr-set720 and parp-1C03256mutants and co-enriched with PARP-1 and H4K20me1 were mainly upregulated (n = 101, 72%) (Fig.3E). Intriguingly, under wild-type conditions, these genes displayed expression levels approximately 40% higher than the average and demonstrated increased of RNA-Polymerase II occupancy both at their promoter regions and gene bodies compared to other genes (Supplemental Fig.S6), indicating their high activity in wild type context. Interestingly, among the genes co-enriched with PARP-1 and H4K20me1, those only differentially expressed in parp-1C03256 mutants were mostly upregulated (n = 94, 85%), while those only differentially expressed in pr-set720 mutants were mainly downregulated (n = 191, 67%), indicating that PARP-1 and H4K20me1 also have distinct functions in gene regulation at co-enriched genes (Fig.3E). Gene ontology analysis of differentially expressed genes bound by PARP-1 and H4K20me1 showed that metabolic genes were upregulated while genes involved in neuron development and morphogenesis were downregulated in both pr-set720and parp-1C03256 mutants (Fig.3F). Finally, to extend the generalizability of our observations beyond Drosophila, we compared the distribution of PARP1 and H4K20me1 in Human K562 cells. Strikingly, we observed a correlation in their distribution, suggesting that the interplay between PARP-1 and H4K20me1 is not limited to fruit flies (Supplemental Fig.S7). Overall, our results indicate H4K20me1 may be required for PARP-1 binding to preferentially repress metabolic genes and activate genes involved in neuron development at co-enriched genes.

Revised version: Next, we examined the relationship between H4K20me1 and PARP-1 during transcriptional regulation. Given that flies that lack PR-SET7, the sole methylase of H4K20, and PARP-1 undergo developmental arrest during the larval-to-pupal transition (14, 15), we hypothesized that PR-SET7/H4K20me1 and PARP-1 may regulate similar gene expression programs during development. To validate our hypothesis, we initially confirmed that the pr-set720mutant not only eliminated PR-SET7 RNA and protein but also abrogated H4K20me1 modification (Supplemental Fig.S3). Interestingly, in the absence of PARP-1, neither PR-SET7 RNA nor protein levels were affected (Supplemental Fig.S4-5), indicating that PARP-1 is not directly implicated in the regulation of pr-set7. This finding contrasts with recent evidence demonstrating PARP1-induced degradation of PR-SET7/SET8 in human cells (16). Subsequently, we integrated PARP-1 and H4K20me1 ChIP-seq data in third-instar larvae and RNA-seq of pr-set720 and parp-1C03256 mutant third-instar larvae. Our data showed that the differentially expressed genes (DEGs) in pr-set720 and parp-1C03256transcriptome were highly correlated (Pearson, r = 0.79) (Fig.3A and B). However, despite the global reduction of H4K20me1, and PARP-1 RNA in pr-set720and parp-1C03256 mutants, respectively (14, 15), there was a limited alteration in gene expression at genes where PARP-1 was co-localized with H4K20me1 (Fig.3C and D). Differentially expressed genes in both pr-set720 and parp-1C03256mutants and co-enriched with PARP-1 and H4K20me1 were mainly upregulated (n = 101, 72%) (Fig.3E). Intriguingly, under wild-type conditions, these genes displayed expression levels approximately 40% higher than the average and demonstrated increased of RNA-Polymerase II occupancy both at their promoter regions and gene bodies compared to other genes (Supplemental Fig.S6), indicating their high activity in wild type context. Interestingly, among the genes co-enriched with PARP-1 and H4K20me1, those only differentially expressed in parp-1C03256 mutants were mostly upregulated (n = 94, 85%), while those only differentially expressed in pr-set720 mutants were mainly downregulated (n = 191, 67%), indicating that PARP-1 and H4K20me1 also have distinct functions in gene regulation at co-enriched genes (Fig.3E). Gene ontology analysis of differentially expressed genes bound by PARP-1 and H4K20me1 showed that metabolic genes were upregulated while genes involved in neuron development and morphogenesis were downregulated in both pr-set720and parp-1C03256 mutants (Fig.3F). Finally, to extend the generalizability of our observations beyond Drosophila, we compared the distribution of PARP1 and H4K20me1 in Human K562 cells. Strikingly, we observed a correlation in their distribution, suggesting that the interplay between PARP-1 and H4K20me1 is not limited to fruit flies (Supplemental Fig.S7). Overall, our results indicate H4K20me1 may be required for PARP-1 binding to preferentially repress metabolic genes and activate genes involved in neuron development at co-enriched genes.

PARP-1 and H4K20me1 are required for the repression of metabolic genes and activation of developmental genes at co-enriched genes.

Earlier version: (A) Scatterplot plot showing correlation of differentially expressed genes (DEGs) in parp-1C03256and pr-set720 (Pearson’s r = 0.79). (B) Heatmap showing the normalized read counts of DEGs in both parp-1C03256 and pr-set720. Normalized read counts are shown as row z-scores. (C-D) Dot plots showing transcriptional changes of genes co-enriched with PARP-1 and H4K20me1 in parp-1C03256and pr-set720 compared to WT at promoters (C) and gene bodies (D). (E) Summary of DEGs in parp-1C03256 and pr-set720 and both mutants that were co-enriched with PARP-1 and H4K20me1. (F) Gene ontology of upregulated (left) and downregulated (right) DEGs in both parp-1C03256and pr-set720 mutants that were co-enriched with PARP-1 and H4K20me1.

Revised version: (A) Scatterplot plot showing correlation of differentially expressed genes (DEGs) in parp-1C03256and pr-set720 (Pearson’s r = 0.79). (B) Heatmap showing the normalized read counts of DEGs in both parp-1C03256 and pr-set720. Normalized read counts are shown as row z-scores. (C-D) Dot plots showing transcriptional changes of genes co-enriched with PARP-1 and H4K20me1 in parp-1C03256 and pr-set720compared to WT at promoters (C) and gene bodies (D). (E) Summary of DEGs in parp-1C03256 and pr-set720 and both mutants that were co-enriched with PARP-1 and H4K20me1. (F) Gene ontology of upregulated (left) and downregulated (right) DEGs in both parp-1C03256 and pr-set720 mutants that were co-enriched with PARP-1 and H4K20me1.

PARP-1 and PR-SET7/H4K20me1 are required for the optimal expression of heat shock genes during heat stress

Earlier version: Next, we examined PARP-1 and H4K20me1 controlled gene expression programs during dynamic transcriptional changes. During heat shock, PARP-1 spreading from the promoter to the gene bodies of hsp70 facilitates nucleosomal loss which leads to transcriptional activation (16, 17). Since H4K20me1 is mostly enriched in gene bodies (Fig.1H), we hypothesized that H4K20me1 may be required for PARP-1 spread at hsp70 during heat shock. First, we examined the expression of heat shock genes before and after 30 minutes heat shock in WT, parp-1C03256 and pr-set720 third-instar larvae. The expression of heat shock genes (hsp22, hsp23, hsp68, hsp83) including hsp70 were significantly reduced in parp-1C03256 and pr-set720 animals after heat shock (Fig.4A). ChIP-seq results revealed that PARP-1 highly occupied the promoters of hsp23, hsp68, hsp70, and hsp83, as well as in the gene body of hsp22, prior to heat shock (Fig. 4B and C). Analysis of H4K20me1 enrichment showed low levels across the gene bodies of hsp22, hsp68, and hsp70 (group A) but high levels at gene bodies of hsp23 and hsp83 (group B) before heat shock (Fig.4B and C). Based on these findings, we hypothesized that an increase in H4K20me1 enrichment in the gene bodies of group A genes would facilitate PARP-1 binding after heat shock, while H4K20me1 enrichment remains unchanged in group B genes thereby facilitating PARP-1 binding and spreading. To our surprise, H4K20me1 levels increased at the gene bodies of group A genes (low H4K20me1 before heat shock) (Fig.4D) but decreased significantly at group B genes (high H4K20me1 before heat shock) after heat shock (Fig.4E). These findings suggest that dynamic changes in H4K20me1 enrichment may regulate the expression of heat shock genes during heat shock and indicate that H4K20me1 may not be necessary for PARP-1 binding and PARP-1-mediated activation of group B genes during heat shock.

Revised version: Next, we examined PARP-1 and H4K20me1 controlled gene expression programs during dynamic transcriptional changes. During heat shock, PARP-1 spreading from the promoter to the gene bodies of hsp70 facilitates nucleosomal loss which leads to transcriptional activation (17, 18). Since H4K20me1 is mostly enriched in gene bodies (Fig.1H), we hypothesized that H4K20me1 may be required for PARP-1 spread at hsp70 during heat shock. First, we examined the expression of heat shock genes before and after 30 minutes heat shock in WT, parp-1C03256 and pr-set720third-instar larvae. The expression of heat shock genes (hsp22, hsp23, hsp68, hsp83) including hsp70 were significantly reduced in parp-1C03256 and pr-set720 animals after heat shock (Fig.4A). ChIP-seq results revealed that PARP-1 highly occupied the promoters of hsp23, hsp68, hsp70, and hsp83, as well as in the gene body of hsp22, prior to heat shock (Fig. 4B and C). Analysis of H4K20me1 enrichment showed low levels across the gene bodies of hsp22, hsp68, and hsp70 (group A) but high levels at gene bodies of hsp23 and hsp83 (group B) before heat shock (Fig.4B and C). Based on these findings, we hypothesized that an increase in H4K20me1 enrichment in the gene bodies of group A genes would facilitate PARP-1 binding after heat shock, while H4K20me1 enrichment remains unchanged in group B genes thereby facilitating PARP-1 binding and spreading. To our surprise, H4K20me1 levels increased at the gene bodies of group A genes (low H4K20me1 before heat shock) (Fig.4D) but decreased significantly at group B genes (high H4K20me1 before heat shock) after heat shock (Fig.4E). These findings suggest that dynamic changes in H4K20me1 enrichment may regulate the expression of heat shock genes during heat shock and indicate that H4K20me1 may not be necessary for PARP-1 binding and PARP-1-mediated activation of group B genes during heat shock.

Dynamic H4K20me1 enrichment regulates the expression of heat-shock genes during heat shock.

Earlier version: (A) Expression of hsp22, hsp23, hsp68, hsp70 and hsp83 in WT, parp-1C03256and pr-set720 third-instar larvae before and after 30 mins heat shock. Data shown are from 3-5 biological replicates. ***P < 0.001, **P < 0.01, *P < 0.05 (Unpaired t-test; Two-tailed). Integrative Genome Viewer (IGV) tracks showing normalized ChIP-seq tracks of PARP-1 and H4K20me1 in third-instar larvae at (B) hsp22, hsp68, hsp68, hsp70 and (C) hsp22, hsp83 before heat shock. Mononucleosome ChIP-qPCR in WT showing enrichment of H4K20me1 at (D) hsp70, hsp22, hsp68 and (E) hsp23 and hsp83 before heat shock (NHS), after 15 mins heat shock and 30 mins heat shock. Primers used spanned the hsp70 locus and the gene bodies of hsp22, hsp23, hsp68, hsp70 and hsp83. Data shown are from 3 biological replicates.

Revised version: (A) Expression of hsp22, hsp23, hsp68, hsp70 and hsp83 in WT, parp-1C03256 and pr-set720third-instar larvae before and after 30 mins heat shock. Data shown are from 3-5 biological replicates. ***P < 0.001, **P < 0.01, *P < 0.05 (Unpaired t-test; Two-tailed). Integrative Genome Viewer (IGV) tracks showing normalized ChIP-seq tracks of PARP-1 and H4K20me1 in third-instar larvae at (B) hsp22, hsp68, hsp68, hsp70 and (C) hsp22, hsp83 before heat shock. Mononucleosome ChIP-qPCR in WT showing enrichment of H4K20me1 at (D) hsp70, hsp22, hsp68 and (E) hsp23 and hsp83 before heat shock (NHS), after 15 mins heat shock and 30 mins heat shock. Primers used spanned the hsp70 locus and the gene bodies of hsp22, hsp23, hsp68, hsp70 and hsp83. Data shown are from 3 biological replicates.

Mutating PR-SET7 methyl transferase disrupts PARP-1 binding to chromatin and diminishes poly(ADP)-ribosylation levels during heat shock stress

Earlier version: Our results showed that prior to heat shock, there was no effect on PARP-1 binding in both WT and pr-set720 background animals to chromatin (Fig.5A). However, after heat shock, we observed that PARP-1 binding was concentrated at specific loci in WT animals, while it was diminished from chromatin in pr-set720animals (Fig.5A). These findings suggest that PR-SET7/H4K20me1 plays a crucial role in mediating PARP1 binding to chromatin during heat stress, thus, controls PARP1-dependent processes in chromatin. Previous studies showed that heat shock results in pADPr accumulation at the heat shock-induced puffs during PARP-1-dependent transcriptional activation (5, 7, 11, 16–18). Therefore, we tested if heat-shock treatment alters pADPr accumulation in pr-set720 animals. As expected, the pADPr level increased by 26.1 times after heat-shock treatment in wild type animals, while in pr-set720 mutants the level of pADPr increased only 2.3 folds (Fig.5B). This observation implies that PR-SET7/H4K20me1 plays a key role in regulating PARP-1 upon environmental stresses such as heat shock.

Revised version: Our results showed that prior to heat shock, there was no effect on PARP-1 binding in both WT and pr-set720 background animals to chromatin (Fig.5A). However, after heat shock, we observed that PARP-1 binding was concentrated at specific loci in WT animals, while it was diminished from chromatin in pr-set720animals (Fig.5A). These findings suggest that PR-SET7/H4K20me1 plays a crucial role in mediating PARP1 binding to chromatin during heat stress, thus, controls PARP1-dependent processes in chromatin. Previous studies showed that heat shock results in pADPr accumulation at the heat shock-induced puffs during PARP-1-dependent transcriptional activation (5, 7, 11, 17–19). Therefore, we tested if heat-shock treatment alters pADPr accumulation in pr-set720animals. As expected, the pADPr level increased by 26.1 times after heat-shock treatment in wild type animals, while in pr-set720 mutants the level of pADPr increased only 2.3 folds (Fig.5B). This observation implies that PR-SET7/H4K20me1 plays a key role in regulating PARP-1 upon environmental stresses such as heat shock.

Mono-methylated histones controls PARP-1 binding along gene body to regulate transcription.

(A) PR-SET7/H4K20me1 is required for holding PARP-1 in chromatin during heat shock (HS). PARP-1 (green) protein recruitment to hsp70 locus (Arrows) in salivary gland polytene nuclei in wild type (WT) and pr-set720 mutant third instar larvae. A single salivary gland polytenized nucleus of wandering third instar larvae is shown for each genotype. Wild type genotype - UASt::PARP-1-EYFP, GAL4[Mz1087.hx]; pr-set7 mutant genotype – UASt::PARP-1-EYFP, GAL4Mz1087.hx; pr-set720. (B) Equal amounts of lysates from the wild-type expressing PARP-1-EYFP (WT) and pr-set720mutants expressing PARP-1-EYFP (pr-set720) grown at 22°C or heat shocked at 37°C for one hour at the third-instar larvae stage were subjected to immunoblot analysis using mouse anti-pADPr (10H), anti-H4K20me1 and anti-histone H4 (loading control) antibodies. (C) Model of PARP-1 regulation by histone modifications. PARP-1 binds to a nucleosome that carries the H2A variant (H2Av) at the promoter region. Upon developmental triggers or heat shock-induced phosphorylation of H2Av, PARP-1 is activated (11). Activated PARP-1 fosters a transcription start site (TSS) that is more accessible, thereby enabling the binding of RNA-Polymerase II (Pol II) and initiation of transcription (26). Following this, the distribution of PARP-1 is further enhanced by active mono-methylated forms of H4K20, H3K4, H3K36, or H3K9. Each of these histone modifications can interchangeably facilitate this function. The spreading of PARP-1 to the gene body contributes to the loosening of chromatin in this region (7, 16, 17, 49). Consequently, this facilitates the transition of Pol II into a productive elongation phase, leading to the generation of a mature transcript.

Discussion

This study illuminates the intricate interactions between PARP-1 and diverse histone modifications, with a particular emphasis on the influence of histone marks on the chromatin binding propensity of PARP-1. We discovered an elevated binding affinity of PARP-1 towards specific mono-methylated active histone modifications such as H4K20me1, H3K4me1, H3K36me1, H3K9me1 and H3K27me1 in vitro (Fig.1B and C). Intriguingly, our results also suggest that the repressive histone modifications, H3K9me2/3 may potentially hinder PARP-1 binding (Fig.1F). The association between PARP-1 peaks and distinct histone marks in Drosophila third-instar larvae, as revealed through ChIP-seq data, corroborates the significance of these histone marks in PARP-1 binding. This is predominantly apparent for H4K20me1, H3K4me1, H3K36me1, H3K9me1 and H3K27me1, while H3K9me2/3 demonstrated a lesser correlation (Fig.1G), further validating the repressive role of these marks.

Earlier version: We additionally examined the potential functional consequences of PARP-1’s interaction with these histone marks, particularly H4K20me1. Mono-methylation of histone H4 on lysine 20 (H4K20me1), catalyzed by PR-SET7, is evolutionarily conserved from yeast to humans (19, 20). H4K20me1, similar to PARP-1, is enriched at highly active genes (21–27). Also, H4K20me1 is critical for transcriptional regulation, maintenance of genome integrity, cell cycle regulation and double-strand break (DSB) DNA damage repair, all biological processes in which PARP-1 plays a role (3, 19, 28–32). Hence, histone modifications such as H4K20me1 may act to regulate the activity of key chromatin remodelers such as PARP-1.

Revised version: We additionally examined the potential functional consequences of PARP-1’s interaction with these histone marks, particularly H4K20me1. Mono-methylation of histone H4 on lysine 20 (H4K20me1), catalyzed by PR-SET7, is evolutionarily conserved from yeast to humans (20, 21). H4K20me1, similar to PARP-1, is enriched at highly active genes (22–28). Also, H4K20me1 is critical for transcriptional regulation, maintenance of genome integrity, cell cycle regulation and double-strand break (DSB) DNA damage repair, all biological processes in which PARP-1 plays a role (3, 20, 29–33). Hence, histone modifications such as H4K20me1 may act to regulate the activity of key chromatin remodelers such as PARP-1.

Earlier version: A recent study demonstrated that in human cells overexpressing PARP-1, PR-SET7/SET8 is degraded (33). This implies that the absence of PARP-1 might lead to increased levels of PR-SET7. However, in our study involving parp-1 mutant in Drosophila third-instar larvae, we observed a slightly different scenario: we detected a minor but not significant reduction in both PR-SET7 RNA and protein levels (Supplemental Fig.S4 and S5). This outcome stands in stark contrast to the previous study’s findings. The discrepancy could be due to the distinct experimental approaches used: the previous research focused on mammalian cells and in vitro experiments, whereas our study examined the functions of PARP-1 in whole Drosophila third-instar larvae during development. Consequently, while PARP-1 may cooperate with PR-SET7 in the context of Drosophila development, it could exhibit antagonistic roles against PR-SET7 in specific cell lines and under certain biological or developmental conditions.

Revised version: A recent study demonstrated that in human cells overexpressing PARP-1, PR-SET7/SET8 is degraded, whereas depletion of PARP-1 leads to an increase in PR-SET7/SET8 levels (16). However, in our study involving parp-1 mutant in Drosophila third-instar larvae revealed a nuanced scenario: we detected a minor but not significant reduction in both PR-SET7 RNA and protein levels (Supplemental Fig.S4 and S5). This outcome stands in stark contrast to the previous study’s findings. The discrepancy could be due to the distinct experimental approaches used: the previous research focused on mammalian cells and in vitro experiments, whereas our study examined the functions of PARP-1 in whole Drosophila third-instar larvae during development. Consequently, while PARP-1 may cooperate with PR-SET7 in the context of Drosophila development, it could exhibit antagonistic roles against PR-SET7 in specific cell lines and under certain biological or developmental conditions.

Our transcriptomic analysis from pr-set720 and parp-1C03256 mutant third-instar larvae implies that PR-SET7/H4K20me1 might be crucial for PARP-1’s role in transcriptional regulation. Remarkably, differentially expressed genes (DEGs) in pr-set720and parp-1C03256 mutants exhibited a high degree of correlation (Fig.3A), although we noticed minimal alteration in gene expression where PARP-1 was co-localized with H4K20me1 across the genome (Fig.3C and D). This suggests that H4K20me1 might only be essential for PARP-1-mediated binding and gene regulation at a select group of genes, specifically during metabolic gene repression and developmental gene activation. Our results also propose that PARP-1 and PR-SET7/H4K20me1 might have separate roles in gene regulation at co-enriched genes.

Earlier version: Notably, genes co-enriched with PARP-1 and H4K20me1, and are upregulated in both parp-1C03256 and pr-set720 mutants, are predominantly metabolic genes exhibiting high expression levels under wild-type conditions and a high occupancy of RNA-polymerase II both at their promoter region and gene body (Supplemental Fig.S6). In our previous studies, we discovered that PARP-1 plays a crucial role in repressing highly active metabolic genes during the development of Drosophila by binding directly to their loci (34, 35). Also, PARP-1 is required for maintaining optimum glucose and ATP levels at the third-instar larval stage (34). During Drosophila development, repression of metabolic genes is crucial for larval to pupal transition (36, 37). This repression is linked to the reduced energy requirements as the organism prepares for its sedentary pupal stage (36, 38). Notably, we observed that PARP-1 shows a high affinity for binding to the gene bodies of these metabolic genes (34).

Revised version: Notably, genes co-enriched with PARP-1 and H4K20me1, and are upregulated in both parp-1C03256 and pr-set720mutants, are predominantly metabolic genes exhibiting high expression levels under wild-type conditions and a high occupancy of RNA-polymerase II both at their promoter region and gene body (Supplemental Fig.S6). In our previous studies, we discovered that PARP-1 plays a crucial role in repressing highly active metabolic genes during the development of Drosophila by binding directly to their loci (34, 35). Also, PARP-1 is required for maintaining optimum glucose and ATP levels at the third-instar larval stage (34). During Drosophila development, repression of metabolic genes is crucial for larval to pupal transition (36, 37). This repression is linked to the reduced energy requirements as the organism prepares for its sedentary pupal stage (36, 38). Remarkably, our observations indicate a notable affinity of PARP-1 for binding to the gene bodies of these metabolic genes (34), suggesting a direct involvement of PARP1 in their regulation. Nonetheless, it remains plausible that certain genes may be indirectly regulated by PARP1 through intermediary transcription factors.

Our data indicates that in both parp-1 and pr-set7 mutant animals, there was a preferential repression of metabolic genes at sites where PARP-1 and H4K20me1 are co-bound (Fig.3E), while these metabolic genes are highly active during third-instar larval stage (Supplemental Fig.S6). Thus, we propose that the presence of H4K20me1 may be essential for the binding of PARP-1 at these gene bodies, contributing to their repression. Importantly, this mechanism of gene repression has broader developmental implications. As earlier stated, mutant animals lacking functional PARP-1 and PR-SET7 undergo developmental arrest during larval to pupal transition. This arrest could be directly linked to the disruption of the normal metabolic gene repression during development. Without the repressive action of PARP-1 and PR-SET7, key metabolic processes might remain unchecked, leading to metabolic imbalances that are incompatible with the normal progression to the pupal stage.

Earlier version: In this study, we have demonstrated that PARP-1 is essential for the activation of hsp22, hsp23, and hsp83 during heat shock (Fig.4A). Importantly, the expression of these heat shock genes significantly decreased in pr-set720 mutants following heat shock (Fig.4A), which shows a role for PR-SET7/H4K20me1 in the heat shock response. We have also observed that H4K20me1 regulation is dynamic in the gene bodies of PARP-1- and PR-SET7-dependent heat shock genes. Specifically, at genes that show low or no H4K20me1 enrichment prior to heat shock, H4K20me1 levels increase post-heat shock (Fig.4B and D). Conversely, at genes with high levels of H4K20me1, the level of H4K20me1 decreases following heat shock (Fig.4C and D). Previous studies have shown that H4K20me1 occupancy can either inhibit or facilitate transcriptional elongation, depending on the biological context (25, 27, 39–41). As such, we propose that for group A genes, which show a significant increase in H4K20me1 enrichment after heat shock, H4K20me1 may facilitate the spread of PARP-1 in their gene bodies, thus enhancing transcriptional activation.

Revised version: In this study, we have demonstrated that PARP-1 is essential for the activation of hsp22, hsp23, and hsp83 during heat shock (Fig.4A). Importantly, the expression of these heat shock genes significantly decreased in pr-set720 mutants following heat shock (Fig.4A), which shows a role for PR-SET7/H4K20me1 in the heat shock response. We have also observed that H4K20me1 regulation is dynamic in the gene bodies of PARP-1- and PR-SET7-dependent heat shock genes. Specifically, at genes that show low or no H4K20me1 enrichment prior to heat shock, H4K20me1 levels increase post-heat shock (Fig.4B and D). Conversely, at genes with high levels of H4K20me1, the level of H4K20me1 decreases following heat shock (Fig.4C and D). Previous studies have shown that H4K20me1 occupancy can either inhibit or facilitate transcriptional elongation, depending on the biological context (26, 28, 39–41). As such, we propose that for group A genes, which show a significant increase in H4K20me1 enrichment after heat shock, H4K20me1 may facilitate the spread of PARP-1 in their gene bodies, thus enhancing transcriptional activation. In contrast, the removal of H4K20me1 at group B genes, which exhibit a significant reduction of H4K20me1 at their gene bodies, may enable transcriptional elongation. Thus, an unidentified H4K20me1 demethylase and PR-SET7’s non-catalytic functions may be required to facilitate transcriptional activation of group B heat shock genes during heat stress.

Removed in the revised version: In contrast, the removal of H4K20me1 at group B genes, which exhibit a significant reduction of H4K20me1 at their gene bodies, may enable transcriptional elongation. Thus, an unidentified H4K20me1 demethylase and PR-SET7’s non-catalytic functions may be required to facilitate transcriptional activation of group B heat shock genes during heat stress.

Earlier version: Another plausible explanation could be that the recruitment of PARP-1 to group B genes loci promotes H4 dissociation and then leads to a reduction of H4K20me1. However, our findings suggest an alternative interpretation: the decrease in H4K20me1 at group B genes during heat shock does not seem to impede PARP-1’s role in transcriptional activation, (Fig.4A, C and E). Rather than disrupting PARP-1 function, we propose that this reduction in H4K20me1 may signify a regulatory shift in chromatin structure, priming these genes for transcriptional activation during heat shock, with PARP-1 playing an independent facilitating role. Moreover, existing studies have highlighted the dual role of H4K20me1, acting as a promoter of transcription elongation in certain contexts and as a repressor in others (13, 25, 39, 40, 42–46). The elevated enrichment of H4K20me1 in group B genes under normal conditions may indicate a repressive state that requires alleviation for transcriptional activation. Additionally, we cannot discount the possibility of unique regulatory functions associated with PR-SET7, extending beyond its recognized role as a histone methylase. Non-catalytic activities and potential interactions with non-histone substrates might contribute to the nuanced control exerted by PR-SET7 on group B genes during heat stress (47, 48). Furthermore, our exploration of pr-set720 and parp-1C03256 mutants reveals distinct roles for PARP-1 and H4K20me1 in modulating gene expression (Fig 3E). This reinforces the notion that the interplay between PR-SET7 and PARP-1 involves a multifaceted regulatory mechanism. Understanding the intricate relationship between these molecular players is crucial for elucidating the complexities of gene expression modulation under heat stress conditions.

Revised version: Another plausible explanation could be that the recruitment of PARP-1 to group B genes loci promotes H4 dissociation and then leads to a reduction of H4K20me1. However, our findings suggest an alternative interpretation: the decrease in H4K20me1 at group B genes during heat shock does not seem to impede PARP-1’s role in transcriptional activation, (Fig.4A, C and E). Rather than disrupting PARP-1 function, we propose that this reduction in H4K20me1 may signify a regulatory shift in chromatin structure, priming these genes for transcriptional activation during heat shock, with PARP-1 playing an independent facilitating role. Moreover, existing studies have highlighted the dual role of H4K20me1, acting as a promoter of transcription elongation in certain contexts and as a repressor in others (13, 26, 39, 40, 42–46). The elevated enrichment of H4K20me1 in group B genes under normal conditions may indicate a repressive state that requires alleviation for transcriptional activation. Additionally, we cannot discount the possibility of unique regulatory functions associated with PR-SET7, extending beyond its recognized role as a histone methylase. Non-catalytic activities and potential interactions with non-histone substrates might contribute to the nuanced control exerted by PR-SET7 on group B genes during heat stress (47, 48). Furthermore, our exploration of pr-set720and parp-1C03256 mutants reveals distinct roles for PARP-1 and H4K20me1 in modulating gene expression (Fig 3E). This reinforces the notion that the interplay between PR-SET7 and PARP-1 involves a multifaceted regulatory mechanism. Understanding the intricate relationship between these molecular players is crucial for elucidating the complexities of gene expression modulation under heat stress conditions.

Notably, Drosophila embryo expressing an unmodifiable H4K20 (H4K20A) and a catalytically-inactive trr mutant (monomethylase of H3K4) were still able to mature into adult flies (14, 49), demonstrating that H3K4me1 and H4K20me1 are dispensable for transcription during development. This finding contrasts our hypothesis that H4K20me1 or H3K4me1 might modulate PARP-1 binding at PARP-1-dependent genes during development. Importantly, our data also revealed that PARP-1 binds to H3K36me1 and H3K9me1. Consequently, we suggest that H3K36me1 and H3K9me1 could potentially substitute the role of H3K4me1 and H4K20me1 in controlling PARP-1 binding at PARP-1 dependent genes (Fig.5C).

Earlier version: Finaly, highly transcribed genes have been reported to present a high turnover of mono-methylated modifications, maintaining a state of low methylation (50). Then, our findings suggest that PARP-1 might actively participate in the turnover process to uphold an active chromatin environment. The proposed mechanism unfolds as follows: 1) PARP-1 selectively binds to mono-methylated active histone marks associated with highly transcribed genes. 2) Upon activation, PARP-1 undergoes automodification and is subsequently released from chromatin, facilitating the reassembly of nucleosomes carrying the mono-methylated marks. 3) The enzymatic action of Poly(ADP)-ribose glycohydrolase (PARG) cleaves pADPr, allowing for the restoration of PARP-1’s binding affinity to mono-methylated active histone marks. This proposed hypothesis aligns cohesively with existing research conducted across various model organisms, including mice, Drosophila, and Humans (7, 23, 29, 51–53). Notably, previous studies have consistently demonstrated that PARP-1 predominantly associates with highly expressed genes and plays a crucial role in mediating nucleosome dynamics and assembly. Thus, our proposed model provides a molecular framework that may contribute to understanding the relationship between PARP-1 and the epigenetic regulation of gene expression. Further experimental validation is warranted to elucidate the precise details of this proposed mechanism and its implications in the broader context of chromatin dynamics and transcriptional control.

Revised version: Finaly, highly transcribed genes have been reported to present a high turnover of mono-methylated modifications, maintaining a state of low methylation (50). Moreover, our previous study revealed that PARP1 preferentially binds to highly active genes (34). Consequently, our findings suggest an active involvement of PARP-1 in the turnover process to maintain an active chromatin environment. This proposed mechanism unfolds in the following steps: 1) PARP-1 selectively binds to mono-methylated active histone marks associated with highly transcribed genes. 2) Upon activation, PARP-1 undergoes automodification and subsequently disengages from chromatin, facilitating the reassembly of nucleosomes carrying the mono-methylated marks. 3) The enzymatic action of Poly(ADP)-ribose glycohydrolase (PARG) cleaves pADPr, restoring PARP-1’s binding affinity to mono-methylated active histone marks. This proposed hypothesis is consistent with existing research conducted across various model organisms, including mice, Drosophila, and Humans (7, 24, 30, 51–53). Notably, previous studies have consistently demonstrated that PARP-1 predominantly associates with highly expressed genes and plays a crucial role in mediating nucleosome dynamics and assembly. Thus, our proposed model provides a molecular framework that may contribute to understanding the relationship between PARP-1 and the epigenetic regulation of gene expression.

Materials and methods

Drosophila strains and genetics

Earlier version: Flies were cultured on standard cornmeal-molasses-agar media at 25°C. The transgenic stock pP{w1, UAST::PARP1-EYFP} has been previously described (52). The parp-1C03256 hypomorph mutant was generated in a single pBac-element mutagenesis screen (54). Parp-1C03256 significantly lowers the level of PARP-1 RNA and protein level (15) but also significantly diminishes the level of pADPr (11). The pr-set720null mutant was generously provided by Dr. Ruth Steward (14). The pr-set720 null mutant was validated in (14) and we confirmed that this mutant abolishes PR-SET7 RNA and protein level but also leads to the absence of H4K20me1 (Supplemental Fig.S3). Wandering third-instar larvae was used for all experiments. w1118 strains were used as controls for ChIP-seq and RNA-seq and termed WT as appropriate. We used TM6B and TM3-GFP to isolate parp-1C03256 and pr-set720 homozygous mutants respectively.

Revised version: Flies were cultured on standard cornmeal-molasses-agar media at 25°C. The transgenic stock pP{w1, UAST::PARP1-EYFP} has been previously described (52). The parp-1C03256hypomorph mutant was generated in a single pBac-element mutagenesis screen (54). Parp-1C03256 significantly lowers the level of PARP-1 RNA and protein level (15) but also significantly diminishes the level of pADPr (11). The pr-set720 null mutant was generously provided by Dr. Ruth Steward (14). The pr-set720 null mutant was validated in (14) and we confirmed that this mutant abolishes PR-SET7 RNA and protein level but also leads to the absence of H4K20me1 (Supplemental Fig.S3). Wandering third-instar larvae was used for all experiments. w1118 strains were used as controls for ChIP-seq and RNA-seq and termed WT as appropriate. We used TM6B and TM3-GFP to isolate parp-1C03256 and pr-set720 homozygous mutants respectively.

Histone peptide array

The Modified Histone peptide array (Active Motif) was blocked by overnight incubation in the tween 20 containing tris buffered saline (TTBS) buffer (10 mm Tris/HCl pH 8.3, 0.05% Tween-20 and 150 mm NaCl) containing 5% non-fat dried milk at 4°C. The membrane was then washed once with the TTBS buffer and incubated with 4.0μgm PARP-1 (Trevigen) in PARP binding buffer (10 mm Tris-HCl, pH 8, 140 mm NaCl, 3 mm DTT, and 0.1% Triton X-100) at room temperature for 1 h. The membrane was washed in the TTBS buffer and incubated with anti-PARP antibody (SERUTEC at 1:500 dilution) for 1 h at room temperature in blocking buffer (5% milk in TTBS). The unbound antibody was washed three times with TTBS, and the membrane was incubated with horseradish peroxidase-conjugated anti-mouse antibody (Sigma, 1:2500) in TTBS for 1 h at room temperature. Finally, the membrane was submerged in ECL developing solution (GE Healthcare) and the image was captured on X-ray film. Typical exposure times were 0.5–2 min. The images were analyzed using an in-house program (Array Analyze, available at https://www.activemotif.com/catalog/668/modified-histone-peptide-array).

Drosophila salivary gland polytene chromosome immunostaining

Preparation and immunostaining of polytene chromosome squashes were performed exactly as described (55). The primary antibody used was anti-GFP (Living Colors, #JL-8, 1:100), H3K4me1 (Abcam, 1:50) and H4K20me1 (Abcam, 1:100) and the secondary antibody used was goat anti-mouse Alexa-488 (Molecular Probes, 1:1500) and goat-anti-rabbit Alexa-568 (Molecular Probes, 1:1500). DNA was stained with TOTO3 (1:3000). Slides were mounted in Vectashield (Vector Laboratories, Burlingame, CA).

Chromatin immunoprecipitation and sequencing

PARP-1-YFP ChIP-seq has been previously described (35). Briefly, 75 third-instar larvae were collected, washed with 1ml 1X PBS, and homogenized in lysis buffer (200ul 1X protease inhibitor cocktail, 250ul PMSF, 800ul 1X PBS, 1ul Tween 20). Crosslinking was achieved by adding 244.5ul of 11% formaldehyde for a final concentration of 1.8% and quenched with 500mM glycine. The pellet was resuspended in sonication buffer (0.5% SDS, 20mM Tris pH 8.0, 2mM EDTA, 0.5mM EGTA, 0.5mM PMSF, 1X protease inhibitor cocktail) and sonicated to fragment chromatin. The supernatant was then collected after pelleting. This was pre-cleared, incubated with anti-GFP antibody (TP-401, Torrey Pines Biolabs), and immunoprecipitated chromatin was collected using Protein A agarose beads. Sequential washing was performed using low salt buffer, high salt buffer, LiCl wash, and TE buffer washers. Bound chromatin was eluted using 250ul ChIP elution buffer (1% SDS, 100mM NaHCO3) and reverse-crosslinked. The eluates were treated with RNase A and proteinase K, and DNA was extracted via phenol-chloroform extraction and ethanol precipitation. Sequencing was carried out at Novogene.

ChIP-seq analysis

The quality of FASTQ files (raw reads) were checked using FastQC (version. 0.11.9) and adapters were removed with fastp (56). Trimmed FASTQ files were aligned to the Drosophila genome (dm6) using Bowtie2 to generate bam files (57). Unmapped and low-quality reads were discarded from bam files (<=20 mapQuality) using BamTools (58). Duplicate reads were identified and removed from mapped reads using Picard MarkDuplicates (http://broadinstitute.github.io/picard/). Deeptools MultiBamSummary was used to determine reproducibility of ChIP-seq reads. MACS2 was used to call peaks against pooled Input/control using default settings except narrowPeaks were called for PARP-1 and broadPeaks were called for H3K4me1 (gapped), H3K9me1, H4K20me1, H3K36me1, H3K9me2, H3K9me3. Peaks were annotated to genomic features with ChIPseeker (59). Pairwise correlation of peaks was determined using Intervene (60). MACS2 bedGraph pileups were used to generate normalized coverage of ChIP-seq signals using Deeptools bigWigCompare by computing the ratio of the signals (IP vs Control/Input) using a 50bp bin size. Deeptools multiBigwigSummary was used to determine genome-wide signal correlation using a 10kb bin size. Deeptools plotHeatmap was used to create gene-centric enrichment profiles using scaled region mode. Gene clusters were determined via K means function (n=3) using Deeptools suite (61).

To compare H4K20me1 and PARP-1 occupancy in Human K562 cells (Supplemental Fig.S7), CrossMap was used to convert H4K20me1 bigwig from Hg19 to Hg38 to match PARP-1 bigwig. Deeptools were used to generate and plot enrichment of PARP-1 and H4K20me1 at PARP-1 gene clusters.

Mononucleosome chromatin immunoprecipitation

Removed in the revised version: We collected 60 WT wandering Drosophila third-instar larvae and rinsed them in 1X PBS twice. For heat shock experiments, larvae were heat-shocked in a 1.5ml DNA LoBind tube for 15mins and 30mins in a water bath at 37°C. Larvae were then homogenized using a pellet pestle homogenizer in 500ul ice-cold buffer A1 (60 mM KCl, 15 mM NaCl, 4 mM MgCl2, 15 mM 4-(2-Hydroxyethyl)piperazine-1-ethanesulfonic acid (HEPES) (pH 7.6), 0.5% Triton X-100, 0.5 mM DTT, 1×EDTA (Ethylene Diamine Tetraacetic Acid)-free protease inhibitor cocktail (Roche 04693132001). The larvae were then crosslinked by adding formaldehyde to a 1.8% concentration and incubating for 5mins at RT on a rotator. Crosslinking was stopped by adding glycine (final concentration was 0.125 M) at room temperature for 5 min on a rotator. The mixture was centrifuged at 2000g for 5 min at 4°C, and the supernatant was discarded. The pellet was washed as follows: once with 500μl A1 buffer, and once with 500μl A2 buffer (140 mM NaCl, 15 mM HEPES (pH 7.6), 1 mM EDTA, 0.5 mM ethylene glycol tetraacetic acid (EGTA), 1% Triton X-100, 0.1% sodium deoxycholate, 0.5 mM DTT, 1×EDTA-free protease inhibitor cocktail (Roche 04693132001)). For each wash, the tube was shaken for 1 min and centrifuged as before, and the supernatant was discarded. The pellet was resuspended in 500μl A2 buffer + 0.1% sodium dodecyl sulfate (SDS) and incubated on a rotating wheel at 4°C for 10 min. Then the mixture was centrifuged at 16 000g for 5 min at 4°C and the supernatant was discarded. The pellet was washed with 500μl MNase digestion buffer (10 mM Tris–HCl (pH 7.5), 15 mM NaCl, 60 mM KCl, 1 mM CaCl2, 0.15 mM spermine, 0.5 mM spermidine, 1×EDTA-free protease inhibitor cocktail) and centrifuged at 16 000g for 10 min at 4°C. The nuclei were resuspended in 500μl MNase digestion buffer and 400 U MNase (Worthington, LS004797) was added and incubated at 37°C for 20 min. The MNase digestion was terminated on ice by adding EDTA to a final concentration of 10 mM and kept on ice for 10 min.

Earlier version: The mixture was centrifuged at 16 000g for 10 min at 4°C. The supernatant was discarded. The pellet was resuspended in 500μl A3 buffer (140 mM NaCl, 15 mM HEPES (pH 7.6), 1 mM EDTA, 0.5 mM EGTA, 1% Triton X-100, 0.1% sodium deoxycholate, 0.1% SDS, 1×EDTA-free protease inhibitor cocktail) and incubated for 30 min on a rotator at 4°C. The mixture was then centrifuged at 16 000g for 15 min at 4°C. The supernatant containing mononucleosomal fragments was transferred to a new DNA LoBind tube, then 20ul of mixture was used to check fragment size prior to immunoprecipitation. The mixture was then made up to 2ml with IP buffer (0.5% SDS, 20mM Tris, pH 8.0, 2mM EDTA, 0.5mM EGTA, 0.5mM PMSF, protease inhibitor cocktail). The mononucleosomal chromatin was pre-cleared and 800ul was incubated with 5ul anti-H4K20me1 antibody (abcam, 9051) or 5ul Rabbit IgG control at 4°C overnight and 80ul was kept as input. The immunoprecipitated chromatin was then collected with pre-washed Protein A agarose beads for 2 hours. The beads were sequentially washed with the following buffers: 1 low salt buffer wash (0.1% SDS, 1% Triton X-100, 2mM EDTA, 20mM Tris-HCL pH 8.0, 150mM NaCl), 3 high salt buffer washes (0.1% SDS, 1% Triton X-100, 2mM EDTA, 20mM Tris-HCL pH 8.0, 500mM NaCl), 1 LiCL wash (2mM EDTA, 20mM Tris-HCl pH 8.0, 0.25M LiCl, 1% NP-40) and 2 TE buffer washers before elution. Bound chromatin on beads was eluted twice at room temperature using 250ul of freshly prepared ChIP elution buffer (1% SDS, 100mM NaHCO3) for 15 minutes and reverse-crosslinked overnight. The eluates were then treated with RNase A and proteinase K prior to DNA extraction via phenol-chloroform extraction and ethanol precipitation.

Revised version: We collected 60 WT wandering Drosophila third-instar larvae and rinsed them in 1X PBS twice. For heat shock experiments, larvae were heat-shocked in a 1.5ml DNA LoBind tube for 15mins and 30mins in a water bath at 37°C. Larvae were then homogenized using a pellet pestle homogenizer in 500ul ice-cold buffer A1 (60 mM KCl, 15 mM NaCl, 4 mM MgCl2, 15 mM 4-(2-Hydroxyethyl)piperazine-1-ethanesulfonic acid (HEPES) (pH 7.6), 0.5% Triton X-100, 0.5 mM DTT, 1×EDTA (Ethylene Diamine Tetraacetic Acid)-free protease inhibitor cocktail (Roche 04693132001). The larvae were then crosslinked by adding formaldehyde to a 1.8% concentration and incubating for 5mins at RT on a rotator. Crosslinking was stopped by adding glycine (final concentration was 0.125 M) at room temperature for 5 min on a rotator. The mixture was centrifuged at 2000g for 5 min at 4°C, and the supernatant was discarded. The pellet was washed as follows: once with 500μl A1 buffer, and once with 500μl A2 buffer (140 mM NaCl, 15 mM HEPES (pH 7.6), 1 mM EDTA, 0.5 mM ethylene glycol tetraacetic acid (EGTA), 1% Triton X-100, 0.1% sodium deoxycholate, 0.5 mM DTT, 1×EDTA-free protease inhibitor cocktail (Roche 04693132001)). For each wash, the tube was shaken for 1 min and centrifuged as before, and the supernatant was discarded. The pellet was resuspended in 500μl A2 buffer + 0.1% sodium dodecyl sulfate (SDS) and incubated on a rotating wheel at 4°C for 10 min. Then the mixture was centrifuged at 16 000g for 5 min at 4°C and the supernatant was discarded. The pellet was washed with 500μl MNase digestion buffer (10 mM Tris–HCl (pH 7.5), 15 mM NaCl, 60 mM KCl, 1 mM CaCl2, 0.15 mM spermine, 0.5 mM spermidine, 1×EDTA-free protease inhibitor cocktail) and centrifuged at 16 000g for 10 min at 4°C. The nuclei were resuspended in 500μl MNase digestion buffer and 400 U MNase (Worthington, LS004797) was added and incubated at 37°C for 20 min. The MNase digestion was terminated on ice by adding EDTA to a final concentration of 10 mM and kept on ice for 10 min. The mixture was centrifuged at 16 000g for 10 min at 4°C. The supernatant was discarded. The pellet was resuspended in 500μl A3 buffer (140 mM NaCl, 15 mM HEPES (pH 7.6), 1 mM EDTA, 0.5 mM EGTA, 1% Triton X-100, 0.1% sodium deoxycholate, 0.1% SDS, 1×EDTA-free protease inhibitor cocktail) and incubated for 30 min on a rotator at 4°C. The mixture was then centrifuged at 16 000g for 15 min at 4°C. The supernatant containing mononucleosomal fragments was transferred to a new DNA LoBind tube, then 20ul of mixture was used to check fragment size prior to immunoprecipitation. The mixture was then made up to 2ml with IP buffer (0.5% SDS, 20mM Tris, pH 8.0, 2mM EDTA, 0.5mM EGTA, 0.5mM PMSF, protease inhibitor cocktail). The mononucleosomal chromatin was pre-cleared and 800ul was incubated with 5ul anti-H4K20me1 antibody (abcam, 9051) or 5ul Rabbit IgG control at 4°C overnight and 80ul was kept as input. The immunoprecipitated chromatin was then collected with pre-washed Protein A agarose beads for 2 hours. The beads were sequentially washed with the following buffers: 1 low salt buffer wash (0.1% SDS, 1% Triton X-100, 2mM EDTA, 20mM Tris-HCL pH 8.0, 150mM NaCl), 3 high salt buffer washes (0.1% SDS, 1% Triton X-100, 2mM EDTA, 20mM Tris-HCL pH 8.0, 500mM NaCl), 1 LiCL wash (2mM EDTA, 20mM Tris-HCl pH 8.0, 0.25M LiCl, 1% NP-40) and 2 TE buffer washers before elution. Bound chromatin on beads was eluted twice at room temperature using 250ul of freshly prepared ChIP elution buffer (1% SDS, 100mM NaHCO3) for 15 minutes and reverse-crosslinked overnight. The eluates were then treated with RNase A and proteinase K prior to DNA extraction via phenol-chloroform extraction and ethanol precipitation.

RNA sequencing

RNA was isolated from 10 wandering third-instar larvae (3 biological replicates per genotype; WT, pr-set720 and parp-1C03256) using RNeasy lipid tissue mini kit (Qiagen). RNA samples were flash-frozen in liquid nitrogen and sent to Novogene for library preparation and sequencing. mRNA was purified from total RNA via poly-T oligo beads. Libraries were prepared using Ultra II RNA library kit (NEB) and samples were sequenced on the NovaSeq 6000 platform (Illumina) at Novogene.

RNA-seq analysis

Paired-end reads were quality-checked using FastQC and reads were mapped to the Drosophila genome (dm6) using RNA STAR (62). Reads per annotated gene was counted using featureCounts (63). Differential expression analysis was performed with DESeq2 (64), with Log2 fold change of at least 1 (absolute) considered significant (FDR < 0.05).

Quantitative RT-PCR

For heat shock experiments, 10 third-instar larvae were heat-shocked in 1.5ml tube in a water bath at 37°C for 30 min. Total RNA was then isolated for three biological replicate using an RNeasy Lipid Tissue Mini kit (Qiagen). DNA contamination was removed using TURBO DNA-free (Ambion). RNA concentration was measured using an Agilent 2100 BioAnalyzer (Agilent Technologies) in combination with an RNA 6000 Nano LabChip (Agilent Technologies). Real-time PCR assays were run using StepOnePlus Real-Time PCR System (Applied Biosystems).

Earlier version: Primers used for hsp70 ChIP-qPCR were gotten from (18, 65). Primers used for hsp22, hsp23, hsp68 and hsp83 qPCR were gotten from (64). Primer sequences are listed in data file S2.

Revised version: Primers used for hsp70 ChIP-qPCR were gotten from (19, 65). Primers used for hsp22, hsp23, hsp68 and hsp83 qPCR were gotten from (64). Primer sequences are listed in data file S2.

pr-set720 mutants validation using qRT-PCR

Total RNA was isolated from 10 third-instar larvae for three biological replicates using an RNeasy Lipid Tissue Mini kit (Qiagen). DNA contamination was removed using TURBO DNA-free (Ambion). RNA concentration was measured using an Agilent 2100 BioAnalyzer (Agilent Technologies) in combination with an RNA 6000 Nano LabChip (Agilent Technologies). Real-time PCR assays were run using StepOnePlus Real-Time PCR System (Applied Biosystems). The following primers were used: pr-set7 (Forward) CGCACAATAGGAGTTCCC (Reverse) CCTCATCGTCCAGTTTCAG and rp49 (Forward) GCTAAGCTGTCGCACAAAT (Reverse) GAACTTCTTGAATCCGGTGGG.

PARP-1-EYFP re-localization in chromatin during heat shock

Larvae were grown at 20°C. Heat shock was performed by placing third instar wandering larvae individually in a 1.5ml Eppendorf tube with a small hole in the cap to let the larvae breath. The tube was placed in a Marry bath at 37°C for 30 minutes. After heat shock, the larvae were briefly washed in 1X PBS then immediately dissected in 1X supplemented Grace’s insect media (Gibco). Salivary glands were mounted in 1X supplemented Grace’s insect media (Gibco) containing Draq-5 DNA dye (Life technologies) at a 1/5000 dilution.

Western Blot Analysis

The following antibodies were used for immunoblotting assays: anti-PR-SET7 (Rabbit, 1:1000, Novus Biologicals, 44710002), anti-H4K20me1 (Rabbit, 1:1000, Abcam, ab9051), anti-H4 (Mouse, 1:1000, Abcam, ab31830), anti-pADPr (Mouse monoclonal, 1:500, 10H - sc-56198, Santa Cruz), anti-H3 (Rabbit polyclonal, 1/1000, FL-136 sc-10809 Santa Cruz).

Gene ontology (GO)

Gene ontology terms were determined using g:profiler (FDR<0.05) (66).

Statistical analysis

Results were analyzed using the indicated statistical test in GraphPad Prism (9.4.0). Statistical significance of ChIP-seq peaks were determined using MACS2. Q-value cutoffs for RNA-seq and GO analysis were determined with DESeq2 and G:profiler respectively.

Supporting information